Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Sreekar Guggilam"
COLLABORATORS = ""
```

---

# Lab 2 Introduction


In this lab you will use parse a vcf file to extracts parts of it and load it into a database. 
The fields you will be parsing from the vcf file are: 
```
CHROM
POS	
ID	
REF	
ALT	
QUAL	
FILTER
```
and from the INFO column, the following fields:
```
1000g2015aug_all
ExAC_ALL
FATHMM_pred
LRT_pred
MetaLR_pred
MetaSVM_pred
MutationAssessor_pred
MutationTaster_pred
PROVEAN_pred
Polyphen2_HDIV_pred
Polyphen2_HVAR_pred
SIFT_pred
fathmm-MKL_coding_pred.
```
The fields:
```
FATHMM_pred
LRT_pred
MetaLR_pred
MetaSVM_pred
MutationAssessor_pred
MutationTaster_pred
PROVEAN_pred
Polyphen2_HDIV_pred
Polyphen2_HVAR_pred
SIFT_pred
fathmm-MKL_coding_pred
```
are predictor fields. They use a letter to indicate whether a given variation is harmful or not.

NOTE: use the helper functions in cell2. To use them, you have to run CELL 2! 


```python
## Helper functions

import os
import sqlite3
from sqlite3 import Error
import gzip

def create_connection(db_file, delete_db=False):
    if delete_db and os.path.exists(db_file):
        os.remove(db_file)
    conn = None
    try:
        conn = sqlite3.connect(db_file)
        conn.execute("PRAGMA foreign_keys = 1")
    except Error as e:
        print(e)

    return conn


def create_table(conn, create_table_sql):
    try:
        c = conn.cursor()
        c.execute(create_table_sql)
    except Error as e:
        print(e)


```

```python
def determine_data_type(value):
    """
    The function takes a string input and determines its data type to be either a float, int, or string. 
    """
    # YOUR CODE HERE
    try :
        data = int(value)
        return type(data)
    except ValueError:
        try :
            data = float(value)
            return type(data)
        except:
            return type(value)
        
```

```python
def determine_data_type_of_list(values):
    """
    Function takes a list of strings and determines their data type. 

    """
    # YOUR CODE HERE
    str_count = 0
    float_count = 0
    int_count = 0
    for i in range(len(values)):
        if determine_data_type(values[i]) == str:
            str_count += 1
        elif determine_data_type(values[i]) == float:
            float_count += 1
        else:
            int_count += 1
    if str_count >= 1:
        return str
    elif float_count >=1 :
        return float
    else:
        return int
```

```python
# format_field = "GT:AD:DP:GQ:PGT:PID:PL"
# sample_field = {'XG102': '1/1:0,76:76:99:1|1:48306945_C_G:3353,229,0',
#              'XG103': '1/1:0,52:52:99:.:.:1517,156,0',
#              'XG104': '0/1:34,38:72:99:.:.:938,0,796',
#              'XG202': '1/1:0,76:76:99:1|1:48306945_C_G:3353,229,0',
#              'XG203': '1/1:0,52:52:99:.:.:1517,156,0',
#              'XG204': '0/1:34,38:72:99:.:.:938,0,796',
#              'XG302': '1/1:0,76:76:99:1|1:48306945_C_G:3353,229,0',
#              'XG303': '1/1:0,52:52:99:.:.:1517,156,0',
#              'XG304': '0/1:34,38:72:99:.:.:938,0,796',
#              'XG402': '1/1:0,76:76:99:1|1:48306945_C_G:3353,229,0',
#              'XG403': '1/1:0,52:52:99:.:.:1517,156,0',
#              'XG404': '0/1:34,38:72:99:.:.:938,0,796'}


def format_sample_fields(format_field, sample_field):
    """
    Format the sample fields given the description above. Data is already provided.
    """
    
    # YOUR CODE HERE
    sample_keys_list = [*sample_field]
    format_keys_list = format_field.split(":")
    output  = {}
    for i in sample_keys_list:
        test_list = sample_field[i].split(":")
        output[i]={} 
        for j in format_keys_list:
            output[i][j] = test_list[format_keys_list.index(j)]
    
    return output
```

```python
# header = "CHROM	POS	ID	REF	ALT	QUAL	FILTER	INFO	FORMAT	XG102	XG103	XG104	XG202	XG203	XG204	XG302	XG303	XG304	XG402	XG403	XG404".split('\t')
# line = """4	123416186	.	A	G	23.25	PASS	AC=1;AF=0.167;AN=6;BaseQRankSum=-2.542;ClippingRankSum=0;DP=180;ExcessHet=3.0103;FS=0;MLEAC=1;MLEAF=0.167;MQ=52.77;MQRankSum=-4.631;QD=0.39;ReadPosRankSum=1.45;SOR=0.758;VQSLOD=-8.209;culprit=MQ;ANNOVAR_DATE=2018-04-16;Func.refGene=intergenic;Gene.refGene=IL2\x3bIL21;GeneDetail.refGene=dist\x3d38536\x3bdist\x3d117597;ExonicFunc.refGene=.;AAChange.refGene=.;Func.ensGene=intergenic;Gene.ensGene=ENSG00000109471\x3bENSG00000138684;GeneDetail.ensGene=dist\x3d38306\x3bdist\x3d117597;ExonicFunc.ensGene=.;AAChange.ensGene=.;cytoBand=4q27;gwasCatalog=.;tfbsConsSites=.;wgRna=.;targetScanS=.;Gene_symbol=.;OXPHOS_Complex=.;Ensembl_Gene_ID=.;Ensembl_Protein_ID=.;Uniprot_Name=.;Uniprot_ID=.;NCBI_Gene_ID=.;NCBI_Protein_ID=.;Gene_pos=.;AA_pos=.;AA_sub=.;Codon_sub=.;dbSNP_ID=.;PhyloP_46V=.;PhastCons_46V=.;PhyloP_100V=.;PhastCons_100V=.;SiteVar=.;PolyPhen2_prediction=.;PolyPhen2_score=.;SIFT_prediction=.;SIFT_score=.;FatHmm_prediction=.;FatHmm_score=.;PROVEAN_prediction=.;PROVEAN_score=.;MutAss_prediction=.;MutAss_score=.;EFIN_Swiss_Prot_Score=.;EFIN_Swiss_Prot_Prediction=.;EFIN_HumDiv_Score=.;EFIN_HumDiv_Prediction=.;CADD_score=.;CADD_Phred_score=.;CADD_prediction=.;Carol_prediction=.;Carol_score=.;Condel_score=.;Condel_pred=.;COVEC_WMV=.;COVEC_WMV_prediction=.;PolyPhen2_score_transf=.;PolyPhen2_pred_transf=.;SIFT_score_transf=.;SIFT_pred_transf=.;MutAss_score_transf=.;MutAss_pred_transf=.;Perc_coevo_Sites=.;Mean_MI_score=.;COSMIC_ID=.;Tumor_site=.;Examined_samples=.;Mutation_frequency=.;US=.;Status=.;Associated_disease=.;Presence_in_TD=.;Class_predicted=.;Prob_N=.;Prob_P=.;SIFT_score=.;SIFT_converted_rankscore=.;SIFT_pred=.;Polyphen2_HDIV_score=.;Polyphen2_HDIV_rankscore=.;Polyphen2_HDIV_pred=.;Polyphen2_HVAR_score=.;Polyphen2_HVAR_rankscore=.;Polyphen2_HVAR_pred=.;LRT_score=.;LRT_converted_rankscore=.;LRT_pred=.;MutationTaster_score=.;MutationTaster_converted_rankscore=.;MutationTaster_pred=.;MutationAssessor_score=.;MutationAssessor_score_rankscore=.;MutationAssessor_pred=.;FATHMM_score=.;FATHMM_converted_rankscore=.;FATHMM_pred=.;PROVEAN_score=.;PROVEAN_converted_rankscore=.;PROVEAN_pred=.;VEST3_score=.;VEST3_rankscore=.;MetaSVM_score=.;MetaSVM_rankscore=.;MetaSVM_pred=.;MetaLR_score=.;MetaLR_rankscore=.;MetaLR_pred=.;M-CAP_score=.;M-CAP_rankscore=.;M-CAP_pred=.;CADD_raw=.;CADD_raw_rankscore=.;CADD_phred=.;DANN_score=.;DANN_rankscore=.;fathmm-MKL_coding_score=.;fathmm-MKL_coding_rankscore=.;fathmm-MKL_coding_pred=.;Eigen_coding_or_noncoding=.;Eigen-raw=.;Eigen-PC-raw=.;GenoCanyon_score=.;GenoCanyon_score_rankscore=.;integrated_fitCons_score=.;integrated_fitCons_score_rankscore=.;integrated_confidence_value=.;GERP++_RS=.;GERP++_RS_rankscore=.;phyloP100way_vertebrate=.;phyloP100way_vertebrate_rankscore=.;phyloP20way_mammalian=.;phyloP20way_mammalian_rankscore=.;phastCons100way_vertebrate=.;phastCons100way_vertebrate_rankscore=.;phastCons20way_mammalian=.;phastCons20way_mammalian_rankscore=.;SiPhy_29way_logOdds=.;SiPhy_29way_logOdds_rankscore=.;Interpro_domain=.;GTEx_V6_gene=.;GTEx_V6_tissue=.;esp6500siv2_all=.;esp6500siv2_aa=.;esp6500siv2_ea=.;ExAC_ALL=.;ExAC_AFR=.;ExAC_AMR=.;ExAC_EAS=.;ExAC_FIN=.;ExAC_NFE=.;ExAC_OTH=.;ExAC_SAS=.;ExAC_nontcga_ALL=.;ExAC_nontcga_AFR=.;ExAC_nontcga_AMR=.;ExAC_nontcga_EAS=.;ExAC_nontcga_FIN=.;ExAC_nontcga_NFE=.;ExAC_nontcga_OTH=.;ExAC_nontcga_SAS=.;ExAC_nonpsych_ALL=.;ExAC_nonpsych_AFR=.;ExAC_nonpsych_AMR=.;ExAC_nonpsych_EAS=.;ExAC_nonpsych_FIN=.;ExAC_nonpsych_NFE=.;ExAC_nonpsych_OTH=.;ExAC_nonpsych_SAS=.;1000g2015aug_all=.;1000g2015aug_afr=.;1000g2015aug_amr=.;1000g2015aug_eur=.;1000g2015aug_sas=.;CLNALLELEID=.;CLNDN=.;CLNDISDB=.;CLNREVSTAT=.;CLNSIG=.;dbscSNV_ADA_SCORE=.;dbscSNV_RF_SCORE=.;snp138NonFlagged=.;avsnp150=.;CADD13_RawScore=0.015973;CADD13_PHRED=2.741;Eigen=-0.3239;REVEL=.;MCAP=.;Interpro_domain=.;ICGC_Id=.;ICGC_Occurrence=.;gnomAD_genome_ALL=0.0003;gnomAD_genome_AFR=0.0001;gnomAD_genome_AMR=0;gnomAD_genome_ASJ=0;gnomAD_genome_EAS=0.0007;gnomAD_genome_FIN=0.0009;gnomAD_genome_NFE=0.0002;gnomAD_genome_OTH=0.0011;gerp++gt2=.;cosmic70=.;InterVar_automated=.;PVS1=.;PS1=.;PS2=.;PS3=.;PS4=.;PM1=.;PM2=.;PM3=.;PM4=.;PM5=.;PM6=.;PP1=.;PP2=.;PP3=.;PP4=.;PP5=.;BA1=.;BS1=.;BS2=.;BS3=.;BS4=.;BP1=.;BP2=.;BP3=.;BP4=.;BP5=.;BP6=.;BP7=.;Kaviar_AF=.;Kaviar_AC=.;Kaviar_AN=.;ALLELE_END	GT:AD:DP:GQ:PL	0/1:51,8:59:32:32,0,1388	0/0:47,0:47:99:0,114,1353	0/0:74,0:74:51:0,51,1827	0/1:51,8:59:32:32,0,1388	0/0:47,0:47:99:0,114,1353	0/0:74,0:74:51:0,51,1827	0/1:51,8:59:32:32,0,1388	0/0:47,0:47:99:0,114,1353	0/0:74,0:74:51:0,51,1827	0/1:51,8:59:32:32,0,1388	0/0:47,0:47:99:0,114,1353	0/0:74,0:74:51:0,51,1827"""

def create_dict_from_line(header, line):
    """
    Given the header and a single line, transform them into dictionary as described above. 
    Header and line input are provided in this cell. 
    """
    # YOUR CODE HERE
    line_data = line.split('\t')
    format_line = {}
    output = {}
    sample_field = {}
    for i in header:
        format_line[i] = line_data[header.index(i)]
    sample_keys = header[header.index("FORMAT")+1:]
    line_keys = header[0:header.index("FORMAT")]
    for i in sample_keys:
        sample_field[i] = format_line[i]
    format_field = format_line["FORMAT"]
    for i in line_keys:
        output[i] = format_line[i]
    output["SAMPLE"] = format_sample_fields(format_field, sample_field)
    return output
    
    
    


```

```python
def read_vcf_file(filename):
    """
    See description above
    """
    # YOUR CODE HERE
    ## here the input is the vcf filename(which has header and line within itself) 
    ## and output is the formatted data in the vcf file as a list of dictionary
    ## note we are using the previous fun to create the dictionary whose input is header(is a list already) and line(is a string).
    variant = []
    output = []
    read_data = open(filename,mode = "r").readlines()
#     print(read_data)
    read_line = []
    read_header = []
#     print(read_data)
    for i in range(len(read_data)): 
        if read_data[i][0] != '#':
            read_line.append(read_data[i].strip("\n"))
        else:
            if read_data[i][1] != '#':
                read_header = read_data[i].strip('\n').strip('#').split('\t')
            else:
                continue
            
    for i in range(len(read_line)):
        output.append(create_dict_from_line(read_header, read_line[i]))
#     print(read_header)
    return output
            
    
        
    
        
```

```python
def extract_info_field(data):
    """
    See description in part 6
    """
    # YOUR CODE HERE
    # here input is the complete vcf file data formatted to dictionary
    # and output is just the list of data in the info keys for all the variants
    info = []
    for i in range(len(data)):
        info.append(data[i]["INFO"])
    return info
```

```python
def create_dictionary_of_info_field_values(data):
    """
    See description of part 7
    """
    # YOUR CODE HERE
    # here data is info data as a list of strings(which is the info data)
    # and output is dictionary of the info data line(for a single variant file)
    info_data = data[:]
    info_dict_list = []
    for i in range(len(info_data)):
        info_data[i] = info_data[i].split(';')
        info_data[i].remove('ALLELE_END')  ## removing the ALLELE_END
        info_dict = {}
        for j in range(len(info_data[i])):
            info_data[i][j] = info_data[i][j].split('=')
            if len(info_data[i][j]) == 2:
                info_data[i][j][1] = info_data[i][j][1].replace('\\x3b',',').replace('\\x3d','=')
                info_dict[info_data[i][j][0]] = info_data[i][j][1]

        info_dict_list.append(info_dict)
    final_info_dict = {}
    for i in info_dict.keys():
        final_info_dict[i] = []
    for i in range(len(info_dict_list)):
        for j in info_dict_list[i].keys():
            if info_dict_list[i][j] != '.':
                final_info_dict[j].append(info_dict_list[i][j])
            
        
    
    
    
#     for i in range(len(info_data)):
#         info_dict = {}
#         for j in range(len(info_data[i])):
#             info_dict[info_data[i][j][0]] = [info_data[i][j][1].replace('\\x3b',',').replace('\\x3d','=')]
#             if info_dict[info_data[i][j][0]] == ["."]:
#                 info_dict[info_data[i][j][0]] = []
#     data = extract_info_field(result)
    return final_info_dict
```

```python
def determine_data_type_of_info_fields(data):
    """
    See desription in part 8
    """
    # YOUR CODE HERE
    #  here data is info - dictionary
    keys_list = [*data]
    output = {}
    for i in keys_list:
        output[i] = determine_data_type_of_list(data[i])
    return output
```

```python
def format_data(data, info_field_data_type):
    # YOUR CODE HERE
    ## here data is the dictionary created from the vcf file without formatting the info data,
    ## info_field_data_type is the dictionary of data type of the values in the info(dictionary) with the keys as the info(dictionary)
    ## output is the same data just formating the info field into dictionary using the fun - create_dictionary_of_info_field_values()
    ## and the update the data format of the individual value of the info keys(only) not the whole datafile dictionary.
    ## and also the 'POS' and 'QUAL' are formatted to their respective data type.
    format_data = data
    for i in range(len(data)):
        info_data = data[i]['INFO']
        info_data = info_data.split(';')
        info_data.remove('ALLELE_END')  ## removing the ALLELE_END
        info_dict = {}
        for j in range(len(info_data)):
            info_data[j] = info_data[j].split('=')
            if len(info_data[j]) == 2:
                info_data[j][1] = info_data[j][1].replace('\\x3b',',').replace('\\x3d','=')
                if info_data[j][1] != '.':
                    info_dict[info_data[j][0]] = info_data[j][1]
                    
                else:
                    continue
            
            else:
                continue
        for k in [*info_dict]:
            info_dict[k] = info_field_data_type[k](info_dict[k])
        format_data[i]['INFO'] = info_dict
        format_data[i]['POS'] = int(data[i]['POS'])
        format_data[i]['QUAL'] = float(data[i]['QUAL'])
        
    return format_data


```

# Part 1 (10 pts)

In part1 you will open the gzipped vcf file using the python gzip module and figure out all the possible values for the predictor fields. 

You do not have to unzip the file. Use `with gzip.open(filename,'rt') as fp:` to read one line at a time. 



<!-- grade_id: cell-6220a32f79b3dbc2, points: 0, grade: false, solution: true -->
```python
def get_predictor_values(filename):
    """
    See part 1 description 
    """
    import gzip

   
    keys_save_value = [
        'FATHMM_pred',
        'LRT_pred',
        'MetaLR_pred',
        'MetaSVM_pred',
        'MutationAssessor_pred',
        'MutationTaster_pred',
        'PROVEAN_pred',
        'Polyphen2_HDIV_pred',
        'Polyphen2_HVAR_pred',
        'SIFT_pred',
        'fathmm-MKL_coding_pred',
    ]

    # YOUR CODE HERE
    with gzip.open(filename,'rt') as fp:
        read_data = fp.read()
    read_data = read_data.split('\n')
    read_data.pop()
    
    read_line = []
    read_header = []
    data = []
#     print(read_data)
    for i in range(len(read_data)): 
        if read_data[i][0] != '#':
            read_line.append(read_data[i].strip("\n"))
        else:
            if read_data[i][1] != '#':
                read_header = read_data[i].strip('\n').strip('#').split('\t')
            else:
                continue
            
    for i in range(len(read_line)):
        data.append(create_dict_from_line(read_header, read_line[i]))
        
    info_data = extract_info_field(data)
    info_dict = create_dictionary_of_info_field_values(info_data)
#     info_datatype = determine_data_type_of_info_fields(info_dict)
    
#     data = format_data(data, info_datatype)
    
    info_data_values = {}
    for i in keys_save_value :
        info_data_values[i] = []
        for j in range(len(info_dict[i])):
            if not info_dict[i][j] in info_data_values[i]:
                info_data_values[i].append(info_dict[i][j])
        
    return info_data_values
```

<!-- grade_id: cell-70269f3e534eaa6a, points: 10, grade: true, solution: false -->
```python

## Can take 30 seconds to run!
expected_solution = {'SIFT_pred': ['D', 'T'], 'Polyphen2_HDIV_pred': ['D', 'B', 'P'], 'Polyphen2_HVAR_pred': ['D', 'B', 'P'], 'LRT_pred': ['D', 'N', 'U'], 'MutationTaster_pred': ['D', 'P', 'N', 'A'], 'MutationAssessor_pred': ['H', 'N', 'L', 'M'], 'FATHMM_pred': ['T', 'D'], 'PROVEAN_pred': ['D', 'N'], 'MetaSVM_pred': ['D', 'T'], 'MetaLR_pred': ['D', 'T'], 'fathmm-MKL_coding_pred': ['D', 'N']}
filename = 'test_4families_annovar.vcf.gz'
predictor_values = get_predictor_values(filename)
assert predictor_values == expected_solution


```

```python
[*predictor_values]
```

Output:
```
['FATHMM_pred',
 'LRT_pred',
 'MetaLR_pred',
 'MetaSVM_pred',
 'MutationAssessor_pred',
 'MutationTaster_pred',
 'PROVEAN_pred',
 'Polyphen2_HDIV_pred',
 'Polyphen2_HVAR_pred',
 'SIFT_pred',
 'fathmm-MKL_coding_pred']
```

# Part 2 (No points)

You will now create 13 tables
Tables 1-11 will be for :

FATHMM_pred
LRT_pred 
MetaLR_pred
MetaSVM_pred
MutationAssessor_pred
MutationTaster_pred
PROVEAN_pred
Polyphen2_HDIV_pred
Polyphen2_HVAR_pred
SIFT_pred
fathmm_MKL_coding_pred ## NOTE: I have replaced the dash with an underscore. 

The first column for each of these tables will be the name of the field + ID, e.g., FATHMM_predID. This
column will be of type not null primary key. The second column will be called `prediction` and will be of
type TEXT not null. The prediction values will be values you extracted for each of the fields in the previous step. 
For example, 'FATHMM_pred' has two prediction values 'T' and 'D'. Name the tables after their field name, e.g.,  
call table that will contains values of FATHMM_pred `FATHMM_pred`.  Make sure to sort the values, meaning, D should 
be inserted before T. 


Table 12 will be the Variants table. 
The first column will be VariantID. 

The other columns will be:
CHROM   
POS 
ID  
REF 
ALT 
QUAL    
FILTER
thousandg2015aug_all ##NOTE: a column name cannot start with a number, so you have to rename!
ExAC_ALL


Table 12 will have 11 more columns that relate to each of prediction table. Consider writing a utility function
to fetch the primary key for a given prediction from each of the prediction table. Name each of the 
column should be the name of predictor + ID, e.g., FATHMM_predID. 


You have already deteremined their data type, so use that to set their data type. 
For integer use INTEGER. 
For float use REAL. 
For string use TEXT. 

Table 13  will be called PredictionStats. The first column will be PredictorStatsID INTEGER NOT NULL PRIMARY KEY. 
The second column will be VariantID. The third column will be PredictorName. The fourth column will be PredictorValue. 
This is not a normalized table!

The prediction value will be a float value mapped from the prediction text.  Use the following information to 
assign values: 
REF: https://brb.nci.nih.gov/seqtools/colexpanno.html#dbnsfp

FATHMM_pred
- T = 0
- D = 1

LRT_pred 
- D = 1
- N = 0
- U = 0

MetaLR_pred
- T = 0
- D = 1

MetaSVM_pred
- T = 0
- D = 1

MutationAssessor_pred
- H = 1
- N = 0
- L = 0.25
- M = 0.5

MutationTaster_pred
- D = 1
- P = 0
- N = 0
- A = 1

PROVEAN_pred
- D = 1
- N = 0

Polyphen2_HDIV_pred
- D = 1
- B = 0
- P = 0.5

Polyphen2_HVAR_pred
- D = 1
- B = 0
- P = 0.5

SIFT_pred
- D = 1
- T = 0

fathmm-MKL_coding_pred
- D = 1
- N = 0


```

prediction_mapping = {
    'FATHMM_pred': {'T': 0, 'D': 1},
    'MetaLR_pred': {'T': 0, 'D': 1},
    'MetaSVM_pred': {'T': 0, 'D': 1},
    'SIFT_pred': {'T': 0, 'D': 1},
    'fathmm_MKL_coding_pred': {'D': 1, 'N': 0},
    'LRT_pred': {'U': 0, 'N': 0, 'D': 1},
    'MutationAssessor_pred': {'H': 1, 'N': 0, 'L': 0.25, 'M': 0.5},  
    'MutationTaster_pred': {'D': 1, 'P': 0, 'N': 0, 'A': 1},
    'PROVEAN_pred': {'D': 1, 'N': 0},
    'Polyphen2_HDIV_pred': {'D': 1, 'B': 0, 'P': 0.5},
    'Polyphen2_HVAR_pred': {'D': 1, 'B': 0, 'P': 0.5},
}


```


The idea is that 11 predictors have been used to annotate all the variants in the file. This table combines all that information in one table. 
By grouping the table based on variantid and summing the prediction values mapped from the prediction text, you can find which variants have a consensus on their being deterimental. 

IMPORTANT: instead of fathmm-MKL_coding_pred use fathmm_MKL_coding_pred everywhere. 
HINT: You have to commit the changes you make to the table, otherwise your changes will not be saved. Consider wrapping all changes inside `with conn:`




```python
# Creating lab2.db 
db_file = 'lab2.db'
conn = create_connection(db_file, delete_db=True)
conn.close()

```

<!-- grade_id: cell-d88abb3c867e6ef4, points: 0, grade: false, solution: true -->
```python
def create_tables_1_11(db_file):
    """
    Create tables 1 to 11.
    
    INPUT: db_file -- Name of database file -- eg. lab2.db
    
    """
    # YOUR CODE HERE
    conn = create_connection(db_file)
    with conn:
        for keys,values in predictor_values.items():
            if '-'in keys:
                keys = keys.replace('-','_')
            create_table_sql = 'CREATE TABLE ' + keys + ' ('+ keys +'ID INTEGER NOT NULL PRIMARY KEY, ' + 'prediction TEXT NOT NULL);'
            create_table(conn,create_table_sql)
            values.sort()
            count = 0
            for j in values:
                count += 1
                insert_sql = 'INSERT INTO '+ keys +'(' + keys + 'ID,' + 'prediction) VALUES(?,?)'
                c = conn.cursor()
                c.execute(insert_sql,(count,j))
                

db_file = 'lab2.db'    
create_tables_1_11(db_file)
```

```python
# conn = create_connection(db_file)
# for i in predictor_values.keys():
#     c = conn.execute('select * from ' + i)
#     print(c.fetchall())
```

# Part 3 (10 pts)

Write a function that returns a dictionary that maps for a given predictor its letter prediction to foreign key value. 
Conver the dash in 'fathmm-MKL_coding_pred' to an underscore. 


<!-- grade_id: cell-c1fed27f8a6c9e23, points: 0, grade: false, solution: true -->
```python
def get_predictor_value_to_fk_map(db_file):
    """
    See part 3 description 
    
    INPUT: db_file -- Name of database file -- eg. lab2.db
    
    """
    # YOUR CODE HERE
    conn = create_connection(db_file)
    with conn :
        dic = {}
        for key in predictor_values.keys():
            if '-' in key :
                key = key.replace('-','_')
            c = conn.cursor()
            selectsql = "SELECT * FROM "+ key +";"
            c.execute(selectsql)
            rows = c.fetchall()
            indic = {}
            for i in rows:
                indic.update({i[1] : i[0]})
            dic.update({key : indic})

    return dic
    

```

```python
# db_file = 'lab2.db'
# predictor_fk_map = get_predictor_value_to_fk_map(db_file)
# predictor_fk_map
```

<!-- grade_id: cell-d6d06a3e0545c576, points: 10, grade: true, solution: false -->
```python
expected_solution = {
    'FATHMM_pred': {'D': 1, 'T': 2}, 
    'LRT_pred': {'D': 1, 'N': 2, 'U': 3}, 
    'MetaLR_pred': {'D': 1, 'T': 2}, 
    'MetaSVM_pred': {'D': 1, 'T': 2}, 
    'MutationAssessor_pred': {'H': 1, 'L': 2, 'M': 3, 'N': 4}, 
    'MutationTaster_pred': {'A': 1, 'D': 2, 'N': 3, 'P': 4}, 
    'PROVEAN_pred': {'D': 1, 'N': 2}, 
    'Polyphen2_HDIV_pred': {'B': 1, 'D': 2, 'P': 3}, 
    'Polyphen2_HVAR_pred': {'B': 1, 'D': 2, 'P': 3}, 
    'SIFT_pred': {'D': 1, 'T': 2}, 
    'fathmm_MKL_coding_pred': {'D': 1, 'N': 2}}

db_file = 'lab2.db'
predictor_fk_map = get_predictor_value_to_fk_map(db_file)
assert predictor_fk_map == expected_solution


```

# Part 4 (No Points)
Create table 12 or the variants table. See description above


<!-- grade_id: cell-86e648d05e51d138, points: 0, grade: false, solution: true -->
```python
def create_variants_table(db_file):
    """
    Part 4
    """
    # YOUR CODE HERE
    variant_list = ['CHROM','POS','ID','REF','ALT','QUAL','FILTER','thousandg2015aug_all','ExAC_ALL']
    conn = create_connection(db_file)
    with conn :
        create_table_sql = '''CREATE TABLE Variants ( VariantID INTEGER NOT NULL PRIMARY KEY)'''
        create_table(conn, create_table_sql)
        for i in variant_list:
            if i == 'POS':
                column_type = 'INTEGER'
            elif i == 'QUAL' or i == 'thousandg2015aug_all' or i == 'ExAC_ALL':
                column_type = 'REAL'
            else :
                column_type = 'TEXT'
            c = conn.cursor()
            c.execute("ALTER TABLE Variants ADD COLUMN '{cn}' {ct}".format(cn=i, ct=column_type))
        key_list = [i for i in predictor_values.keys()]
        key_list.sort()
        for key in key_list : 
            if '-' in key :
                key = key.replace('-','_')
            c = conn.cursor()
            c.execute('PRAGMA TABLE_INFO({})'.format(key))
            table_schema = c.fetchall()
            column_names = [tup[1] for tup in table_schema]
            column_types = [tup[2] for tup in table_schema]
            c.execute("ALTER TABLE Variants ADD COLUMN '{cn}' {ct}".format(cn=column_names[0], ct=column_types[0]))
        
        
        

# create table
db_file = 'lab2.db'
create_variants_table(db_file)

```

```python
# conn = create_connection(db_file)
# edit = 'SELECT * VARIANTS;'
# conn.execute(edit)
# # for i in predictor_values.keys():
# #     c = conn.execute('SHOW COLUMNS FROM [table name]' + i)
# #     print(c.fetchall())
```

# Part 5 (No Points)

Create table 13 -- or the prediction stats table. See description above. 


<!-- grade_id: cell-9410edb2ed7c5bb3, points: 0, grade: false, solution: true -->
```python
def create_predictionstats_table(db_file):
    """
    Part 5   
    """
    # YOUR CODE HERE
    conn = create_connection(db_file)
    with conn :
        create_table_sql = '''CREATE TABLE PredictionStats ( PredictorStatsID INTEGER NOT NULL PRIMARY KEY , VariantID INTEGER , PredictorName TEXT, PredictorValue REAL )'''
        create_table(conn, create_table_sql)
    
db_file = 'lab2.db'
create_predictionstats_table(db_file)
```

```python

```

# Part 6 (10 Points)

Write a function to pull the following info fields given the whole info field. 
```
values_to_pull = [
        '1000g2015aug_all',
        'ExAC_ALL',
        'FATHMM_pred',
        'LRT_pred',
        'MetaLR_pred',
        'MetaSVM_pred',
        'MutationAssessor_pred',
        'MutationTaster_pred',
        'PROVEAN_pred',
        'Polyphen2_HDIV_pred',
        'Polyphen2_HVAR_pred',
        'SIFT_pred',
        'fathmm-MKL_coding_pred',
    ]
```


<!-- grade_id: cell-1a912b3aacdae96d, points: 0, grade: false, solution: true -->
```python
def pull_info_values(info):
    """
    See part 6 description
    """
    # YOUR CODE HERE
    values_to_pull = ['1000g2015aug_all','ExAC_ALL','FATHMM_pred','LRT_pred','MetaLR_pred','MetaSVM_pred','MutationAssessor_pred','MutationTaster_pred','PROVEAN_pred','Polyphen2_HDIV_pred','Polyphen2_HVAR_pred','SIFT_pred','fathmm-MKL_coding_pred']    
    dic_info_values = {}
    info_list = info.split(";")
    for j in info_list:
        j.strip()
        if '=' in j :
            kv_list = j.split("=")
            if kv_list[0] in values_to_pull :
                if '1000' in kv_list[0] :
                    kv_list[0] = kv_list[0].replace('1000','thousand')
                elif '-' in kv_list[0] :
                    kv_list[0] = kv_list[0].replace('-','_')

                dic_info_values.update({kv_list[0]:kv_list[1]})
    return dic_info_values

```

<!-- grade_id: cell-daa4d9dd1bb8e921, points: 10, grade: true, solution: false -->
```python
sample_info_input = "AC=2;AF=0.333;AN=6;BaseQRankSum=2.23;ClippingRankSum=0;DP=131;ExcessHet=3.9794;FS=2.831;MLEAC=2;MLEAF=0.333;MQ=60;MQRankSum=0;QD=12.06;ReadPosRankSum=-0.293;SOR=0.592;VQSLOD=21.79;culprit=MQ;DB;POSITIVE_TRAIN_SITE;ANNOVAR_DATE=2018-04-16;Func.refGene=exonic;Gene.refGene=MAST2;GeneDetail.refGene=.;ExonicFunc.refGene=nonsynonymous_SNV;AAChange.refGene=MAST2:NM_015112:exon29:c.G3910A:p.V1304M;Func.ensGene=exonic;Gene.ensGene=ENSG00000086015;GeneDetail.ensGene=.;ExonicFunc.ensGene=nonsynonymous_SNV;AAChange.ensGene=ENSG00000086015:ENST00000361297:exon29:c.G3910A:p.V1304M;cytoBand=1p34.1;gwasCatalog=.;tfbsConsSites=.;wgRna=.;targetScanS=.;Gene_symbol=.;OXPHOS_Complex=.;Ensembl_Gene_ID=.;Ensembl_Protein_ID=.;Uniprot_Name=.;Uniprot_ID=.;NCBI_Gene_ID=.;NCBI_Protein_ID=.;Gene_pos=.;AA_pos=.;AA_sub=.;Codon_sub=.;dbSNP_ID=.;PhyloP_46V=.;PhastCons_46V=.;PhyloP_100V=.;PhastCons_100V=.;SiteVar=.;PolyPhen2_prediction=.;PolyPhen2_score=.;SIFT_prediction=.;SIFT_score=.;FatHmm_prediction=.;FatHmm_score=.;PROVEAN_prediction=.;PROVEAN_score=.;MutAss_prediction=.;MutAss_score=.;EFIN_Swiss_Prot_Score=.;EFIN_Swiss_Prot_Prediction=.;EFIN_HumDiv_Score=.;EFIN_HumDiv_Prediction=.;CADD_score=.;CADD_Phred_score=.;CADD_prediction=.;Carol_prediction=.;Carol_score=.;Condel_score=.;Condel_pred=.;COVEC_WMV=.;COVEC_WMV_prediction=.;PolyPhen2_score_transf=.;PolyPhen2_pred_transf=.;SIFT_score_transf=.;SIFT_pred_transf=.;MutAss_score_transf=.;MutAss_pred_transf=.;Perc_coevo_Sites=.;Mean_MI_score=.;COSMIC_ID=.;Tumor_site=.;Examined_samples=.;Mutation_frequency=.;US=.;Status=.;Associated_disease=.;Presence_in_TD=.;Class_predicted=.;Prob_N=.;Prob_P=.;SIFT_score=0.034;SIFT_converted_rankscore=0.440;SIFT_pred=D;Polyphen2_HDIV_score=0.951;Polyphen2_HDIV_rankscore=0.520;Polyphen2_HDIV_pred=P;Polyphen2_HVAR_score=0.514;Polyphen2_HVAR_rankscore=0.462;Polyphen2_HVAR_pred=P;LRT_score=0.002;LRT_converted_rankscore=0.368;LRT_pred=N;MutationTaster_score=1.000;MutationTaster_converted_rankscore=0.810;MutationTaster_pred=D;MutationAssessor_score=1.67;MutationAssessor_score_rankscore=0.430;MutationAssessor_pred=L;FATHMM_score=1.36;FATHMM_converted_rankscore=0.344;FATHMM_pred=T;PROVEAN_score=-1.4;PROVEAN_converted_rankscore=0.346;PROVEAN_pred=N;VEST3_score=0.158;VEST3_rankscore=0.189;MetaSVM_score=-1.142;MetaSVM_rankscore=0.013;MetaSVM_pred=T;MetaLR_score=0.008;MetaLR_rankscore=0.029;MetaLR_pred=T;M-CAP_score=.;M-CAP_rankscore=.;M-CAP_pred=.;CADD_raw=4.716;CADD_raw_rankscore=0.632;CADD_phred=24.6;DANN_score=0.998;DANN_rankscore=0.927;fathmm-MKL_coding_score=0.900;fathmm-MKL_coding_rankscore=0.506;fathmm-MKL_coding_pred=D;Eigen_coding_or_noncoding=c;Eigen-raw=0.461;Eigen-PC-raw=0.469;GenoCanyon_score=1.000;GenoCanyon_score_rankscore=0.747;integrated_fitCons_score=0.672;integrated_fitCons_score_rankscore=0.522;integrated_confidence_value=0;GERP++_RS=4.22;GERP++_RS_rankscore=0.490;phyloP100way_vertebrate=4.989;phyloP100way_vertebrate_rankscore=0.634;phyloP20way_mammalian=1.047;phyloP20way_mammalian_rankscore=0.674;phastCons100way_vertebrate=1.000;phastCons100way_vertebrate_rankscore=0.715;phastCons20way_mammalian=0.999;phastCons20way_mammalian_rankscore=0.750;SiPhy_29way_logOdds=17.151;SiPhy_29way_logOdds_rankscore=0.866;Interpro_domain=.;GTEx_V6_gene=ENSG00000162415.6;GTEx_V6_tissue=Nerve_Tibial;esp6500siv2_all=0.0560;esp6500siv2_aa=0.0160;esp6500siv2_ea=0.0761;ExAC_ALL=0.0553;ExAC_AFR=0.0140;ExAC_AMR=0.0386;ExAC_EAS=0.0005;ExAC_FIN=0.0798;ExAC_NFE=0.0788;ExAC_OTH=0.0669;ExAC_SAS=0.0145;ExAC_nontcga_ALL=0.0541;ExAC_nontcga_AFR=0.0129;ExAC_nontcga_AMR=0.0379;ExAC_nontcga_EAS=0.0004;ExAC_nontcga_FIN=0.0798;ExAC_nontcga_NFE=0.0802;ExAC_nontcga_OTH=0.0716;ExAC_nontcga_SAS=0.0144;ExAC_nonpsych_ALL=0.0496;ExAC_nonpsych_AFR=0.0140;ExAC_nonpsych_AMR=0.0386;ExAC_nonpsych_EAS=0.0005;ExAC_nonpsych_FIN=0.0763;ExAC_nonpsych_NFE=0.0785;ExAC_nonpsych_OTH=0.0638;ExAC_nonpsych_SAS=0.0145;1000g2015aug_all=0.024361;1000g2015aug_afr=0.0038;1000g2015aug_amr=0.0461;1000g2015aug_eur=0.0795;1000g2015aug_sas=0.0041;CLNALLELEID=.;CLNDN=.;CLNDISDB=.;CLNREVSTAT=.;CLNSIG=.;dbscSNV_ADA_SCORE=.;dbscSNV_RF_SCORE=.;snp138NonFlagged=rs33931638;avsnp150=rs33931638;CADD13_RawScore=4.716301;CADD13_PHRED=24.6;Eigen=0.4614;REVEL=0.098;MCAP=.;Interpro_domain=.;ICGC_Id=.;ICGC_Occurrence=.;gnomAD_genome_ALL=0.0507;gnomAD_genome_AFR=0.0114;gnomAD_genome_AMR=0.0430;gnomAD_genome_ASJ=0.1159;gnomAD_genome_EAS=0;gnomAD_genome_FIN=0.0802;gnomAD_genome_NFE=0.0702;gnomAD_genome_OTH=0.0695;gerp++gt2=4.22;cosmic70=.;InterVar_automated=Benign;PVS1=0;PS1=0;PS2=0;PS3=0;PS4=0;PM1=0;PM2=0;PM3=0;PM4=0;PM5=0;PM6=0;PP1=0;PP2=0;PP3=0;PP4=0;PP5=0;BA1=1;BS1=1;BS2=0;BS3=0;BS4=0;BP1=0;BP2=0;BP3=0;BP4=0;BP5=0;BP6=0;BP7=0;Kaviar_AF=0.0552127;Kaviar_AC=8536;Kaviar_AN=154602;ALLELE_END"

expected_solution = {
    'thousandg2015aug_all': '0.024361', 
    'ExAC_ALL': '0.0553', 
    'SIFT_pred': 'D', 
    'Polyphen2_HDIV_pred': 'P', 
    'Polyphen2_HVAR_pred': 'P', 
    'LRT_pred': 'N',
    'MutationTaster_pred': 'D', 
    'MutationAssessor_pred': 'L', 
    'FATHMM_pred': 'T', 
    'PROVEAN_pred': 'N', 
    'MetaSVM_pred': 'T', 
    'MetaLR_pred': 'T', 
    'fathmm_MKL_coding_pred': 'D'
}


solution  = pull_info_values(sample_info_input)
assert solution == expected_solution

```

# Part 7 (10 points)

Remember that to insert a record in SQLite, you have to use `cur.execute(sql, values)`, where `sql` is the insert statement and `values` is a list/tuple of values that will be substituted into the `sql` string wherever there is a question mark. 

Write a function that takes in as input: CHROM, POS, ID, REF, ALT, QUAL, FILTER, info_values and returns a list with the values in the following order:
```
CHROM 
POS
ID
REF
ALT
QUAL
FILTER
thousandg2015aug_all # 1000 has been replaced by the text thousand
ExAC_ALL
FATHMM_pred
LRT_pred
MetaLR_pred
MetaSVM_pred
MutationAssessor_pred
MutationTaster_pred
PROVEAN_pred
Polyphen2_HDIV_pred
Polyphen2_HVAR_pred
SIFT_pred
fathmm_MKL_coding_pred # note that the dash has been replaced by underscore
```

The info_values dictionary contains the predictor values. Use `None` for any empty/missing value for info fields, thousandg2015aug_all, and ExAC_ALL. 



<!-- grade_id: cell-64622ed1cd1edf93, points: 0, grade: false, solution: true -->
```python
def build_values_list(CHROM, POS, ID, REF, ALT, QUAL, FILTER, info_values):
    """
    See part 7 description 
    
    """
    # YOUR CODE HERE
    build_values = [CHROM, POS, ID, REF, ALT, QUAL, FILTER]
    
    temp = 'thousandg2015aug_all'
    if temp in info_values.keys():
        if info_values[temp] == '' or info_values[temp] == '.':
                info_values[temp] = None 
        build_values.append(info_values[temp])
    else : build_values.append(None)
    temp = 'ExAC_ALL'
    if temp in info_values.keys():
        if info_values[temp] == '' or info_values[temp] == '.':
                info_values[temp] = None 
        build_values.append(info_values[temp])
    else : build_values.append(None) 
        
    key_list = [i for i in predictor_fk_map.keys()]
    key_list.sort()
    for key in key_list:
        if key in info_values:

            if '-' in key :
                key = key.replace('-','_')
            if info_values[key] in predictor_fk_map[key].keys() :
                build_values.append(predictor_fk_map[key][info_values[key]])
            else : build_values.append(None)
        else : build_values.append(None)
            
    return build_values


```

<!-- grade_id: cell-c318c4d870f5531b, points: 10, grade: true, solution: false -->
```python
CHROM, POS, ID, REF, ALT, QUAL, FILTER = (7, 87837848, '.', 'C', 'A', 418.25, 'PASS') 
info_values = {'SIFT_pred': 'D', 'Polyphen2_HDIV_pred': 'D', 'Polyphen2_HVAR_pred': 'D', 'LRT_pred': 'D', 'MutationTaster_pred': 'D', 'MutationAssessor_pred': 'H', 'FATHMM_pred': 'T', 'PROVEAN_pred': 'D', 'MetaSVM_pred': 'D', 'MetaLR_pred': 'D', 'fathmm_MKL_coding_pred': 'D'}

results = build_values_list(CHROM, POS, ID, REF, ALT, QUAL, FILTER, info_values)
expected_results = [7, 87837848, '.', 'C', 'A', 418.25, 'PASS', None, None, 2, 1, 1, 1, 1, 2, 1, 2, 2, 1, 1]
assert results == expected_results

```

# Part 8 (No Points)

Create a function that takes the database `conn` and `values` from the `build_values_list` function to insert a variant record. 

IMPORTANT: Function should return the id of the row inserted, which will be VariantId


<!-- grade_id: cell-72a446c435bc6eb5, points: 0, grade: false, solution: true -->
```python
def insert_variant(conn, values):
    """
    See description Part 8
    """
    # YOUR CODE HERE
    with conn :
        cur = conn.cursor()
        cur.execute('''SELECT VariantID FROM Variants;''')
        ivalues = [len(cur.fetchall()) + 1]
        ivalues.extend(values)
        insertsql = 'INSERT INTO Variants VALUES (?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?, ?)'
        cur = conn.cursor()
        cur.execute(insertsql,ivalues)
        c = cur.lastrowid
    return c
```

```python

```

# Part 9 (No Points)

Create a function that takes the database `conn` and `values` which is the tuple `(VariantId, PredictorName, PredictorValue)` Where `VariantId` will be the return value from the `insert_variant` function,  `PredictorName` is the name of the predictor, and `PredictorValue` is the mapping of the text value to a numeric value. 


<!-- grade_id: cell-6d17ca28ee066a3b, points: 0, grade: false, solution: true -->
```python
def insert_predictionstat(conn, values):
    """
    See description in part 9
    """
    # YOUR CODE HERE
    with conn :
        c= conn.cursor()
        c.execute('''SELECT VariantID FROM PredictionStats;''')
        ivalues = [len(c.fetchall()) + 1]
        ivalues.extend(values)
        insertsql = 'INSERT INTO PredictionStats VALUES (?, ?, ?, ?)'
        c= conn.cursor()
        c.execute(insertsql,ivalues)
    

```

```python

```

# Part 10 (No Points)

Write a function to insert records into both the variants and predictor_stats table. 
Hint:
1) Open connection to database  
2) Read file one line at a time using gzip read
3) Extract CHROM, POS, ID, REF, ALT, QUAL, FILTER, INFO
4) Use the pull_info_values function
5) Use build_values_list 
6) Use insert_variant -- save variant_id
7) Use insert_predictionstat -- insert each predictor at a time and remember to use `prediction_mapping` mapping. 




```

prediction_mapping = {
    'FATHMM_pred': {'T': 0, 'D': 1},
    'MetaLR_pred': {'T': 0, 'D': 1},
    'MetaSVM_pred': {'T': 0, 'D': 1},
    'SIFT_pred': {'T': 0, 'D': 1},
    'fathmm_MKL_coding_pred': {'D': 1, 'N': 0},
    'LRT_pred': {'U': 0, 'N': 0, 'D': 1},
    'MutationAssessor_pred': {'H': 1, 'N': 0, 'L': 0.25, 'M': 0.5},  
    'MutationTaster_pred': {'D': 1, 'P': 0, 'N': 0, 'A': 1},
    'PROVEAN_pred': {'D': 1, 'N': 0},
    'Polyphen2_HDIV_pred': {'D': 1, 'B': 0, 'P': 0.5},
    'Polyphen2_HVAR_pred': {'D': 1, 'B': 0, 'P': 0.5},
}


```


<!-- grade_id: cell-09175a401cb7fde3, points: 0, grade: false, solution: true -->
```python
prediction_mapping = {
    'FATHMM_pred': {'T': 0, 'D': 1},
    'MetaLR_pred': {'T': 0, 'D': 1},
    'MetaSVM_pred': {'T': 0, 'D': 1},
    'SIFT_pred': {'T': 0, 'D': 1},
    'fathmm_MKL_coding_pred': {'D': 1, 'N': 0},
    'LRT_pred': {'U': 0, 'N': 0, 'D': 1},
    'MutationAssessor_pred': {'H': 1, 'N': 0, 'L': 0.25, 'M': 0.5},  
    'MutationTaster_pred': {'D': 1, 'P': 0, 'N': 0, 'A': 1},
    'PROVEAN_pred': {'D': 1, 'N': 0},
    'Polyphen2_HDIV_pred': {'D': 1, 'B': 0, 'P': 0.5},
    'Polyphen2_HVAR_pred': {'D': 1, 'B': 0, 'P': 0.5},
}


def populate_variants_predictorstats_tables(db_file, filename):
    """
    See description in part 10
    """
    # YOUR CODE HERE
    conn = create_connection(db_file)
    with gzip.open(filename,'rt') as fp :
        header=[]
        line=[]
        for l in fp:
            if l.startswith('##'):
                continue # lines starting with double hash are skipped
            else:
                if l.startswith('#'):
                    header.append(l.strip()[1:])
                else:
                    line.append(l.strip())   
        extract_list = ['CHROM','POS','ID','REF','ALT','QUAL','FILTER','INFO']
        variant_list_dict=[]
        for i in line:
            dic = dict(zip(header[0].split("\t"), i.split("\t")))
            variant_list_dict.append(dic)
        for j in variant_list_dict:
            dic_info_values = pull_info_values(j['INFO'])
            build_values = build_values_list(j['CHROM'],j['POS'],j['ID'],j['REF'],j['ALT'],j['QUAL'],j['FILTER'],dic_info_values)
            variantid = insert_variant(conn,build_values)
            for k,v in dic_info_values.items() :
                if k in prediction_mapping.keys():
                    if v in prediction_mapping[k].keys():
                        values = (variantid,k,prediction_mapping[k][v])
                        insert_predictionstat(conn, values)
try:
    conn.close()
except:
    pass

db_file = 'lab2.db'
filename = 'test_4families_annovar.vcf.gz'
populate_variants_predictorstats_tables(db_file, filename)
```

```python

```

# Part 11 (10 Points)

Write a function that returns the total number of variants



<!-- grade_id: cell-1d18f5ef024c5d12, points: 0, grade: false, solution: true -->
```python
def num_of_total_variants(conn):
    # YOUR CODE HERE
    with conn:
        c = conn.cursor()
        c.execute('''SELECT VariantID FROM Variants;''')
        count = len(c.fetchall())
    return count

```

<!-- grade_id: cell-880f6f16bf3fe3e3, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert num_of_total_variants(conn) == 50001
conn.close()
```

# Part 12 (10 Points)
Write a function returns the total number of variant predictions -- the count of the predictiostats table. 

<!-- grade_id: cell-375db6c511630ece, points: 0, grade: false, solution: true -->
```python
def num_of_total_variant_predictions(conn):
    """
    Part 12
    """
    # YOUR CODE HERE
    with conn:
        c = conn.cursor()
        c.execute('''SELECT PredictorStatsID FROM PredictionStats;''')
        count = len(c.fetchall())
    return count


```

<!-- grade_id: cell-4a9db4195a7202c1, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert num_of_total_variant_predictions(conn) == 1324
conn.close()
```

# Part 13 (10 Points)
Return the total number of variant predictions that have value greater than zero. Number of values from the predictionstats table that are greater than 0. 

<!-- grade_id: cell-d511bb5062a366a8, points: 0, grade: false, solution: true -->
```python
def num_of_total_variant_predictions_with_value_gt_zero(conn):
    """
    See part 13 description
    """
    # YOUR CODE HERE
    with conn:
        c = conn.cursor()
        c.execute('''SELECT PredictorStatsID FROM PredictionStats WHERE PredictorValue > 0;''')
        count = len(c.fetchall())
    return count


```

<!-- grade_id: cell-9f84bca2161e571b, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert num_of_total_variant_predictions_with_value_gt_zero(conn) == 219
conn.close()
```

# Part 14 (10 Points)

Write a function that given `CHROM, POS, ID, REF, ALT` returns a variant's info with the following columns (column order is important) :
```
Variants.CHROM,
Variants.POS,
Variants.ID,
Variants.REF,
Variants.ALT,
Variants.QUAL,
Variants.FILTER,
Variants.thousandg2015aug_all,
Variants.ExAC_ALL,
FATHMM_pred.prediction,
LRT_pred.prediction,
MetaLR_pred.prediction,
MetaSVM_pred.prediction,
MutationAssessor_pred.prediction,
MutationTaster_pred.prediction,
PROVEAN_pred.prediction,
Polyphen2_HDIV_pred.prediction,
Polyphen2_HVAR_pred.prediction,
SIFT_pred.prediction,
fathmm_MKL_coding_pred.prediction,
sum(PredictionStats.PredictorValue)
```

For the predictions, return the actual text value. And the last column is the sum of all the mapped prediction scores for a given variant!



<!-- grade_id: cell-e0384800793c5ef7, points: 0, grade: false, solution: true -->
```python
def fetch_variant(conn, CHROM, POS, ID, REF, ALT):
    """
    See Part 14 description
    """
    # YOUR CODE HERE
    with conn:
        c = conn.cursor()
        c.execute('select Variants.CHROM,Variants.POS,Variants.ID,Variants.REF,Variants.ALT,Variants.QUAL,Variants.FILTER,Variants.thousandg2015aug_all,Variants.ExAC_ALL,FATHMM_pred.prediction,LRT_pred.prediction,MetaLR_pred.prediction,MetaSVM_pred.prediction,MutationAssessor_pred.prediction,MutationTaster_pred.prediction,PROVEAN_pred.prediction,Polyphen2_HDIV_pred.prediction,Polyphen2_HVAR_pred.prediction,SIFT_pred.prediction,fathmm_MKL_coding_pred.prediction,sum(PredictionStats.PredictorValue) from Variants LEFT JOIN FATHMM_pred USING(FATHMM_predID) LEFT JOIN LRT_pred USING(LRT_predID) LEFT JOIN MetaLR_pred USING(MetaLR_predID) LEFT JOIN MetaSVM_pred USING(MetaSVM_predID) LEFT JOIN MutationAssessor_pred USING(MutationAssessor_predID) LEFT JOIN MutationTaster_pred USING(MutationTaster_predID) LEFT JOIN PROVEAN_pred USING(PROVEAN_predID) LEFT JOIN Polyphen2_HDIV_pred USING(Polyphen2_HDIV_predID) LEFT JOIN Polyphen2_HVAR_pred USING(Polyphen2_HVAR_predID) LEFT JOIN SIFT_pred USING(SIFT_predID) LEFT JOIN fathmm_MKL_coding_pred USING(fathmm_MKL_coding_predID) LEFT JOIN PredictionStats USING(VariantID) where Variants.CHROM = ?  and Variants.POS = ? and Variants.ID = ? and Variants.REF = ? and Variants.ALT = ? GROUP BY PredictionStats.VariantID ;' , (CHROM, POS,ID, REF, ALT,))
        row = c.fetchall()
        return row[0]
    
```

<!-- grade_id: cell-42ef13abc31ede02, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert fetch_variant(conn, '22', 25599849, 'rs17670506', 'G', 'A') == ('22', 25599849, 'rs17670506', 'G', 'A', 3124.91, 'PASS', 0.0251597, 0.0425, 'D', 'D', 'T', 'T', 'M', 'D', 'D', 'D', 'D', 'D', 'D', 8.5)
conn.close()
```

<!-- grade_id: cell-5fd3bc522d861215, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert fetch_variant(conn, 'X', 2836184, 'rs73632976', 'C', 'T') == ('X', 2836184, 'rs73632976', 'C', 'T', 1892.12, 'PASS', None, 0.0427, 'D', 'U', 'D', 'T', 'M', 'P', 'D', 'P', 'P', 'D', 'D', 6.5)
conn.close()
```

<!-- grade_id: cell-e3f85c8b7beab623, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert fetch_variant(conn, '5', 155935708, 'rs45559835', 'G', 'A') == ('5', 155935708, 'rs45559835', 'G', 'A', 1577.12, 'PASS', 0.0189696, 0.0451, 'D', 'D', 'T', 'T', 'L', 'D', 'D', 'P', 'B', 'T', 'D', 5.75)
conn.close()
```

<!-- grade_id: cell-500f892df9ffe02f, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert fetch_variant(conn, '4', 123416186, '.', 'A', 'G') == ('4', 123416186, '.', 'A', 'G', 23.25, 'PASS', None, None, None, None, None, None, None, None, None, None, None, None, None, None)
conn.close()
```

# Part 15 (10 Points)
Write a function that returns the variant with the highest predictor score sum. 

Return the variant info the following order:
```
Variants.CHROM,
Variants.POS,
Variants.ID,
Variants.REF,
Variants.ALT,
Variants.QUAL,
Variants.FILTER,
Variants.thousandg2015aug_all,
Variants.ExAC_ALL,
FATHMM_pred.prediction,
LRT_pred.prediction,
MetaLR_pred.prediction,
MetaSVM_pred.prediction,
MutationAssessor_pred.prediction,
MutationTaster_pred.prediction,
PROVEAN_pred.prediction,
Polyphen2_HDIV_pred.prediction,
Polyphen2_HVAR_pred.prediction,
SIFT_pred.prediction,
fathmm_MKL_coding_pred.prediction,
sum(PredictionStats.PredictorValue)
        
```

Again, return the predictor text values and the last column is the sum of the prediction values. 

<!-- grade_id: cell-6b1c044f90162322, points: 0, grade: false, solution: true -->
```python
def variant_with_highest_sum_of_predictor_value(conn):
    """
    See part 15 description 
    """
    # YOUR CODE HERE
    with conn:
        c = conn.cursor()
        c.execute('select Variants.CHROM,Variants.POS,Variants.ID,Variants.REF,Variants.ALT,Variants.QUAL,Variants.FILTER,Variants.thousandg2015aug_all,Variants.ExAC_ALL,FATHMM_pred.prediction,LRT_pred.prediction,MetaLR_pred.prediction,MetaSVM_pred.prediction,MutationAssessor_pred.prediction,MutationTaster_pred.prediction,PROVEAN_pred.prediction,Polyphen2_HDIV_pred.prediction,Polyphen2_HVAR_pred.prediction,SIFT_pred.prediction,fathmm_MKL_coding_pred.prediction,sum(PredictionStats.PredictorValue) from Variants LEFT JOIN FATHMM_pred USING(FATHMM_predID) LEFT JOIN LRT_pred USING(LRT_predID) LEFT JOIN MetaLR_pred USING(MetaLR_predID) LEFT JOIN MetaSVM_pred USING(MetaSVM_predID) LEFT JOIN MutationAssessor_pred USING(MutationAssessor_predID) LEFT JOIN MutationTaster_pred USING(MutationTaster_predID) LEFT JOIN PROVEAN_pred USING(PROVEAN_predID) LEFT JOIN Polyphen2_HDIV_pred USING(Polyphen2_HDIV_predID) LEFT JOIN Polyphen2_HVAR_pred USING(Polyphen2_HVAR_predID) LEFT JOIN SIFT_pred USING(SIFT_predID) LEFT JOIN fathmm_MKL_coding_pred USING(fathmm_MKL_coding_predID) LEFT JOIN PredictionStats USING(VariantID) GROUP BY PredictionStats.VariantID HAVING SUM(PredictionStats.PredictorValue) = (SELECT MAX(sum) FROM (SELECT SUM(PredictionStats.PredictorValue) as sum FROM Variants LEFT JOIN PredictionStats USING(VariantID) GROUP BY PredictionStats.VariantID));')
        row = c.fetchall()
        return row[0]


```

<!-- grade_id: cell-9c7365267a16b42f, points: 10, grade: true, solution: false -->
```python
db_file = 'lab2.db'
conn = create_connection(db_file)
assert variant_with_highest_sum_of_predictor_value(conn) == ('7', 87837848, '.', 'C', 'A', 418.25, 'PASS', None, None, 'T', 'D', 'D', 'D', 'H', 'D', 'D', 'D', 'D', 'D', 'D', 10.0)
conn.close()
```

```python

```